# Working with Dictionaries

```python
from nose.tools import assert_equals, assert_true
from ipywidgets import interact
from IPython.display import HTML, display
```

**Exercise:** Create a dictionary variable named "sequences" that has as values sequences1 through sequence3 and keys "seq1", "seq2", and "seq3".

```python
# the "\" character can be used as a line continuation in PPython
sequence1 = \
"""AGTTGTTAGTCTACGTGGACCGACAAGAACAGTTTCGAATCGGAAGCTTGCTTAACGTAGTTCTAACAGT
TTTTTATTAGAGAGCAGATCTCTGATGAACAACCAACGGAAAAAGACGGGTCGACCGTCTTTCAATATGC"""
sequence2= \
"""TGAAACGCGCGAGAAACCGCGTGTCAACTGTTTCACAGTTGGCGAAGAGATTCTCAAAAGGATTGCTTTC
AGGCCAAGGACCCATGAAATTGGTGATGGCTTTTATAGCATTCCTAAGATTTCTAGCCATACC"""
sequence3= \
"""GCAGGAATTTTGGCTAGATGGGGCTCATTCAAGAAGAATGGAGCGATCAAAGTGTTACGGGTCAAGA
AAGAAATCTCAAACATGTTGAACATAATGAACAGGAGGAAAAGATCTGTGACCATGCTCCTCATGCTGCT"""
sequence4= \
"""GCCCACAGCCCTGGCGTTCCATCTGACCACCCGAGGGGGAGAGCCGCACATGATAGTTAGCAAGCAGGAA
AGAGGAAAATCACTTTTGTTTAAGACCTCTGCAGGTGTCAACATGTGCACCCTTATTGCAATGGATTTGG
GAGAGTTATGTGAGGACACAATGACCTACAAATGCCCCCGGATCACTGAGACGGAACCAGATGACGTTGA
CTGTTGGTGCAATGCCACGGAGACATGGGTGACCTATGGAACATGTTCTCAAACTGGTGAACACCGACGA
GACAAACGTTCCGTCGCACTGGCACCACACGTAGGGCTTGGTCTAGAAACAAGAACCGAAACGTGGATGT"""

```

<!-- grade_id: cell-087ca748f1868500, points: 0, grade: false, solution: true -->
```python
sequences = {"seq1":sequence1, "seq2":sequence2, "seq3":sequence3} 
#remember keys have to be immutable
#or, assuming you have some lists...
values = [sequence1, sequence2, sequence3, sequence4]
keys = ["seq1", "seq2", "seq3", "seq4"]
dict1 = {}
for i in range(len(keys)):
    dict1[keys[i]] = values[i]
    
```

<!-- grade_id: cell-9702a4cf0625d955, points: 0, grade: true, solution: false -->
```python
assert_true(type(sequences),dict)
assert_true("seq1" in sequences)
assert_true("seq4" not in sequences)
assert_true(sequences["seq2"]==sequence2)
```

```python
list(zip(keys, values))
#can pass a list of key-value pair tuples to dict e.g dict(zip(keys, values))
```

Output:
```
[('seq1',
  'AGTTGTTAGTCTACGTGGACCGACAAGAACAGTTTCGAATCGGAAGCTTGCTTAACGTAGTTCTAACAGT\nTTTTTATTAGAGAGCAGATCTCTGATGAACAACCAACGGAAAAAGACGGGTCGACCGTCTTTCAATATGC'),
 ('seq2',
  'TGAAACGCGCGAGAAACCGCGTGTCAACTGTTTCACAGTTGGCGAAGAGATTCTCAAAAGGATTGCTTTC\nAGGCCAAGGACCCATGAAATTGGTGATGGCTTTTATAGCATTCCTAAGATTTCTAGCCATACC'),
 ('seq3',
  'GCAGGAATTTTGGCTAGATGGGGCTCATTCAAGAAGAATGGAGCGATCAAAGTGTTACGGGTCAAGA\nAAGAAATCTCAAACATGTTGAACATAATGAACAGGAGGAAAAGATCTGTGACCATGCTCCTCATGCTGCT'),
 ('seq4',
  'GCCCACAGCCCTGGCGTTCCATCTGACCACCCGAGGGGGAGAGCCGCACATGATAGTTAGCAAGCAGGAA\nAGAGGAAAATCACTTTTGTTTAAGACCTCTGCAGGTGTCAACATGTGCACCCTTATTGCAATGGATTTGG\nGAGAGTTATGTGAGGACACAATGACCTACAAATGCCCCCGGATCACTGAGACGGAACCAGATGACGTTGA\nCTGTTGGTGCAATGCCACGGAGACATGGGTGACCTATGGAACATGTTCTCAAACTGGTGAACACCGACGA\nGACAAACGTTCCGTCGCACTGGCACCACACGTAGGGCTTGGTCTAGAAACAAGAACCGAAACGTGGATGT')]
```

```python
{keys[i]: values[i] for i in range(len(keys))}
```

Output:
```
{'seq1': 'AGTTGTTAGTCTACGTGGACCGACAAGAACAGTTTCGAATCGGAAGCTTGCTTAACGTAGTTCTAACAGT\nTTTTTATTAGAGAGCAGATCTCTGATGAACAACCAACGGAAAAAGACGGGTCGACCGTCTTTCAATATGC',
 'seq2': 'TGAAACGCGCGAGAAACCGCGTGTCAACTGTTTCACAGTTGGCGAAGAGATTCTCAAAAGGATTGCTTTC\nAGGCCAAGGACCCATGAAATTGGTGATGGCTTTTATAGCATTCCTAAGATTTCTAGCCATACC',
 'seq3': 'GCAGGAATTTTGGCTAGATGGGGCTCATTCAAGAAGAATGGAGCGATCAAAGTGTTACGGGTCAAGA\nAAGAAATCTCAAACATGTTGAACATAATGAACAGGAGGAAAAGATCTGTGACCATGCTCCTCATGCTGCT',
 'seq4': 'GCCCACAGCCCTGGCGTTCCATCTGACCACCCGAGGGGGAGAGCCGCACATGATAGTTAGCAAGCAGGAA\nAGAGGAAAATCACTTTTGTTTAAGACCTCTGCAGGTGTCAACATGTGCACCCTTATTGCAATGGATTTGG\nGAGAGTTATGTGAGGACACAATGACCTACAAATGCCCCCGGATCACTGAGACGGAACCAGATGACGTTGA\nCTGTTGGTGCAATGCCACGGAGACATGGGTGACCTATGGAACATGTTCTCAAACTGGTGAACACCGACGA\nGACAAACGTTCCGTCGCACTGGCACCACACGTAGGGCTTGGTCTAGAAACAAGAACCGAAACGTGGATGT'}
```

**Exercise:** Modify the function in the cell below, so that it returns the count of how many times the kmer occurs in sequence. Before counting, make sure both `sequence` and `kmer` consists only of valid DNA sequence symbols (i.e. `A`, `C`, `T`, and `G`). If either is not a proper DNA sequence we will have the function fail (raise an Exception)

<!-- grade_id: cell-465b81aa4bb8bd74, points: 0, grade: false, solution: true -->
```python
def kmer_count(sequence="", kmer = "CCGATTCG"):
    """This is an example of defining a function with keyword arguments:
    
    sequence and kmer are variables in the function kmer_present. 
    
    These variables have DEFAULT VALUES ("" and "CCGATTCG" respectively) 
    that will be used if no values are provided explicitly"""
    
    #check validity with sets -- is ACTG a subset of kmer?
    sequence=sequence.replace("\n","").upper()
    kmer=kmer.replace("\n","").upper()
    
    if not set(kmer).issubset(set("ATCG")) or not set(sequence).issubset(set("ATCG")):
        raise ValueError("Invalid DNA")
    
    return sequence.count(kmer)
    
    #see class upload
```

```python
testkmer="CCAT"
testkmer_set=set(kmer)
testkmer_set.issubset(set("ATCG"))
```

Output:
```
False
```

<!-- grade_id: cell-b04563ddd6e6dece, points: 0, grade: true, solution: false -->
```python
assert_equals(kmer_count(sequence4, kmer="CTG"),7)
assert_equals(kmer_count(sequence4, kmer="CACACGTAGGG"),1)
assert_equals(kmer_count(sequence4, kmer="CAGG"), 2)
assert_equals(kmer_count(sequence4, kmer="CAGGTTT"), 0)
```

Output:
```
/opt/conda/lib/python3.5/site-packages/ipykernel_launcher.py:1: DeprecationWarning: Please use assertEqual instead.
  """Entry point for launching an IPython kernel.
/opt/conda/lib/python3.5/site-packages/ipykernel_launcher.py:2: DeprecationWarning: Please use assertEqual instead.
  
/opt/conda/lib/python3.5/site-packages/ipykernel_launcher.py:3: DeprecationWarning: Please use assertEqual instead.
  This is separate from the ipykernel package so we can avoid doing imports until
/opt/conda/lib/python3.5/site-packages/ipykernel_launcher.py:4: DeprecationWarning: Please use assertEqual instead.
  after removing the cwd from sys.path.
```

<!-- grade_id: cell-d67e935ca18fc09d, points: 0, grade: false, solution: true -->
```python
@interact(sequence = sequences, kmer = ""
) #decorators make new things out of functions: create a dropdown with a dictionary input & text area with a string input
def kmer_present(sequence="", kmer = "CCGATTCG"):
    """This is an example of defining a function with keyword arguments:
    
    sequence and kmer are variables in the function kmer_present. 
    
    These variables have DEFAULT VALUES ("" and "CCGATTCG" respectively) that will be used if no values are provided
    explicitly.
    """
    sequence=sequence.upper()
    kmer=kmer.upper()
    
    # Here we are going to create a string with HTML markup. We will use the string format() method to insert
    # our values into the string.
    
    # create a variable to store the length of the sequence
    seq_length = len(sequence) # replace this with actual value
    
    # create a variable to store the answer of whether the kmer is in the sequence
    s = '<h3>KMER Values:</h3><table>\n'
    s += '<tr><td>{0}</td><td>{1}</td></tr>\n'.format("sequence length",seq_length)
    try: # we try to do something, if it works, we do that block. if we get a value error, go to the except block.
        count = kmer_count(sequence, kmer)
        s+= '<tr><td>{0}</td><td>{1}</td></tr>\n'.format("kmer: "+kmer.upper(),count)
        if count >0:
            high_sequence = sequence.replace(kmer, '<font color="red">%s</font>'%kmer)
        else:
            high_sequence = sequence
    except ValueError:
        s+= '<tr><td>{0}</td><td>{1}</td></tr>\n'.format("kmer: "+kmer.upper(),"Not a valid kmer")
        high_sequence = sequence
    s += '<tr><td>{0}</td><td>{1}</td></tr>'.format("sequence ", high_sequence)    
    s += '</table>'

    display(HTML(s))
```

Output:
```
A Jupyter Widget
```

```python

```