# Biomedical Data and Python Types
* What would be the appropriate data type for the following quantities?
    * Patient pulse as recorded at a clinic visit
    * Patient blood pressure 
    * My medical record number 
    * Signal measured in an MRI scanner 
    * Average breaths per minute over an hour
    * The length of a tibia 
    * My weight in kilograms 
    * Days since last menustral period 
    * DNA sequence 



```python
from nose.tools import assert_equal, assert_almost_equal

```

<!-- grade_id: cell-a1b2dbae850acdb8, points: 0, grade: false, solution: true -->
```python
# import things here
import math
import cmath
import random
```

```python
help(random)
```

Output:
```
Help on module random:

NAME
    random - Random variable generators.

MODULE REFERENCE
    https://docs.python.org/3.5/library/random.html
    
    The following documentation is automatically generated from the Python
    source files.  It may be incomplete, incorrect or include features that
    are considered implementation detail and may vary between Python
    implementations.  When in doubt, consult the module reference at the
    location listed above.

DESCRIPTION
        integers
        --------
               uniform within range
    
        sequences
        ---------
               pick random element
               pick random sample
               generate random permutation
    
        distributions on the real line:
        ------------------------------
               uniform
               triangular
               normal (Gaussian)
               lognormal
               negative exponential
               gamma
               beta
               pareto
         
```

### Modify `sqrt1` in order to answer the following questions
#### What is the sqrt of 46.3?

<!-- grade_id: cell-a30df0a6c70ff677, points: 0, grade: false, solution: true -->
```python

def sqrt1(x):
    return math.sqrt(x)
    raise NotImplementedError()

```

<!-- grade_id: cell-bf1080267f921fec, points: 0, grade: true, solution: false -->
```python
assert_almost_equal(sqrt1(46.3),6.8044103344816)
assert_almost_equal(sqrt1(49),7)
assert_almost_equal(sqrt1(525), 22.912878474)

```

#### What is the sqrt of 5-3.4j?

<!-- grade_id: cell-fb1ea205873506c0, points: 0, grade: false, solution: true -->
```python
def sqrt2(x):
    return cmath.sqrt(x)
    raise NotImplementedError()
```

<!-- grade_id: cell-ef37e0b25e844bbf, points: 0, grade: true, solution: false -->
```python
assert_almost_equal(sqrt2(-1),0+1j)
assert_almost_equal(sqrt2(5-3.4j),2.3501581418207147-0.7233555775455076j)
assert_almost_equal(sqrt2(-9),0+3j)
```

### Modify `random_gauss` to return a Gaussian (normal) distributed number with mean `mu` and standard deviation `sigma`
#### Return a floating point number from a Gaussian distribution with $\mu=24$ and $\sigma=2.6$

<!-- grade_id: cell-dbc91f0480157a0a, points: 0, grade: false, solution: true -->
```python
def random_gauss(mu, sigma):
    return random.gauss(mu, sigma)
    raise NotImplementedError()
```

<!-- grade_id: cell-ad25179e230634c2, points: 0, grade: true, solution: false -->
```python
random.seed(100)
assert_almost_equal(random_gauss(24, 2.6),25.746038650049936)

```

### Modify `triangle_area so that it returns the area of a triangle given the length of its three sides a, b, and c using these formulas:
$s=\frac{(a+b+c)}{2}$
$A=\sqrt{s(s-a)(s-b)(s-c)}$


<!-- grade_id: cell-056e9faa8d5a8dac, points: 0, grade: false, solution: true -->
```python
def triangle_area(a, b, c):
    s = .5*(a+b+c)
    A = math.sqrt(s*(s-a)*(s-b)*(s-c))
    return A
    raise NotImplementedError()
```

<!-- grade_id: cell-71a982f8adb8d09c, points: 0, grade: true, solution: false -->
```python
assert_almost_equal(triangle_area(3,4,5),6)
assert_almost_equal(triangle_area(1,1,math.sqrt(2)), 0.5)
assert_almost_equal(triangle_area(10, 15, 20), 72.61843774)

```

### Additional In-class notebooks
* [Who suffered most?](./who_suffered_most.ipynb)
* [blood pressure](./blood_pressure.ipynb)
* [Random arrays of numbers](./RandomArrays.ipynb)
* [DICOM images](./dicom_images.ipynb)
* [DICOM volumes](./DICOMVolumes.ipynb)

```python

```